Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Svitlana Midianko"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    maxvalue = float("-inf")  # setting a current 'max' value to lowest 
    
    # nested for loop which will check susequently _all_ possible combinations of the contiuguous subarray and 
    # uopdate the maxvalue if any of such subarrays are bigger than current maxvalue.
    for i in range(len(A)): 
        runtotal = 0 # updating current running total each time we create a new subarray.
        for j in range(i, len(A)): 
            runtotal += A[j] 
            if runtotal > maxvalue:
                starti = i
                maxvalue = runtotal
                endi = j
    return (starti, endi, float(maxvalue)) # since the task was to output float, we will 'float' maxvalue. 
# initially _print(type(maxvalue))_ = <class 'int'>
    raise NotImplementedError()
A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)
# complexity O(n^2)
```

Output:
```
(2, 2, 4.0)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
My educated guess that yes, in the defined above drute-force approach, one will need to re-execute the whole function with input `B` to find out the solution. This happens because we did not account for a case of _checking_ an extra appended number in the function, thus the only way to figure out the max subarray is to again re-run the function. 
<br>
 **However,** another approach could be to run a `bruteforce_max_subarray(A)` for this sublist - `A[2, len(A)]` (`2` is the index for 4, max subarray found in the previous Qn). In such case, we would avoid rechecking elements `(A[0:2]`) that do not have max subarray for sure. 
This is only valid if the new elements are _appended_ (added linearly _at the end)._

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
`False`. 
<br> 
Since all the numbers, except for one, are the same in `A` and `B`, the newly max subarray of `B` will only be _new_ (different from subarray found in `A`) _if and only if_ the newly added element in `B` changes it.   

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    # specifying condition *:
    if len(x) == 1: 
        return x[0] # if there is only one element in a list, then it is max. 
    
    max_now = mx # since mx = max from previous subarray, we can set it for condtional check and search for
    # larger sum that mx.
    sum_now = 0 # setting accumulator to 0 for the subsequent iteration.
    
    for i in range(len(x)-1,-1,-1): # iterating over all the values to compute new max. len(x) - 1 specifies 
        # the beginning of iteration, -1 the end (which is the first element of the array), -1 - because 
        # we want 'downto'
        sum_now += x[i]
        
        if sum_now > max_now: # checking if the max of each subarray bigger than the current one.
            max_now = sum_now # updating max if so.
            
    return max_now

B = [-2, -3, 4, -1, -2, 1, 8]
incremental_max_subarray(B, 4)
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    max_now = float("-inf") # setting the comparison point to the lowest possible.
    
    # looping and creating all possible subarrays. Updating mx (max_now) when certain
    # subarray in this iteration results in bigger max subarray sum:
    for i in range(len(A)): 
        sum_arr = incremental_max_subarray(A[:i+1], max_now)
        
        if sum_arr>max_now: 
            max_now = sum_arr
            
    return max_now 
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
It is not, since the `max_subarray` function contains for loop (contributing with $O(n)$ complexity) _and_ inside for  loop another function `incremental_max_subarray` that itslef has another for loop. This results in the $O(n^2)$ complexity while divide and conquer has $P(nlogn)$, thus is more efficient. 